rule_files:
    - alerts.yaml

evaluation_interval: 1m

tests:
    # Test that alerts will be triggered.
    - interval: 1m
      input_series:
          - series: 'router_backend_handler_response_duration_seconds_count{response_code="500", namespace="apps", job="router"}'
            values: '0 2+1x6 8x4' # 0 3 4 5 6 7 8 8 8 8 8
          - series: 'router_backend_handler_response_duration_seconds_count{namespace="apps", job="router"}'
            values: '0+1x11' # 1 2 3 4 5 6 7 8 9 10 11

      # Unit test for alerting rules.
      alert_rule_test:
          - eval_time: 11m
            alertname: RouterBackendHandlerResponse5xxErrorRate
            exp_alerts:
                # This alert will trigger when the apps/router error ratio is >0.1 for more than 10 minutes.
                - exp_labels:
                      severity: page
                      job: router

    # Control test, should pass without any alerts triggered.
    - interval: 1m
      input_series:
          - series: 'router_backend_handler_response_duration_seconds_count{namespace="apps", job="router"}'
            values: '0+1x11' # 1 2 3 4 5 6 7 8 9 10 11
          - series: 'router_backend_handler_response_duration_seconds_count{response_code="500", namespace="apps", job="router"}'
            # No alert triggered as only increase in router errors for 4 minutes
            values: '0 2+1x4 6x6' # 0 3 4 5 6 6 6 6 6 6 6

      # Unit test for alerting rules.
      alert_rule_test:
          - eval_time: 11m
            alertname: RouterBackendHandlerResponse5xxErrorRate
            exp_alerts:
                # No alerts triggered.
