rule_files:
  - chat_ai.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # No alert with 5xx errors < 10%
      - series: >-
          http_requests_total{
          job="govuk-chat",
          status="200"
          }
        values: '1000+100x15'
      - series: >-
          http_requests_total{
          job="govuk-chat",
          status="500"
          }
        values: '0+10x15'
    alert_rule_test:
      - alertname: High5xxRate
        eval_time: 15m
        exp_alerts: []

# High5xxRate alert tests
  - interval: 1m
    input_series:
      # No alert with 503 errors > 10%
      - series: >-
          http_requests_total{
          job="govuk-chat",
          status="200"
          }
        values: '1000x15'
      - series: >-
          http_requests_total{
          job="govuk-chat",
          status="503"
          }
        values: '0+100x15'
    alert_rule_test:
      - alertname: High5xxRate
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert with 5xx errors > 10%
      - series: >-
          http_requests_total{
          job="govuk-chat",
          status="200"
          }
        values: '1000x15'
      - series: >-
          http_requests_total{
          job="govuk-chat",
          status="500"
          }
        values: '0+100x15'
    alert_rule_test:
      - alertname: High5xxRate
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: High5xxRate
              severity: page
            exp_annotations:
              summary: Elevated rate of 5xx return codes for Chat AI
              description: >-
                The rate of http 5xx return codes is above 10% of total requests
                for more than 5 minutes.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

# LongRequestDuration alert tests
  - interval: 1m
    input_series:
      # No alert for LongRequestDuration
      - series: >-
          http_request_duration_seconds_sum{
          job="govuk-chat",
          }
        values: '50+1x15'
      - series: >-
          http_request_duration_seconds_count{
          job="govuk-chat",
          }
        values: '20000+1000x15'
    alert_rule_test:
      - alertname: LongRequestDuration
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert for LongRequestDuration
      - series: >-
          http_request_duration_seconds_sum{
          job="govuk-chat"
          }
        values: '0+10x15'
      - series: >-
          http_request_duration_seconds_count{
          job="govuk-chat"
          }
        values: '0+1x15'
    alert_rule_test:
      - alertname: LongRequestDuration
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: LongRequestDuration
              severity: critical
              destination: slack-chat-notifications
            exp_annotations:
              summary: Elevated HTTP request duration for Chat AI
              description: >-
                75th percentile HTTP request duration over 1 second for more than 5 minutes.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

# HighPodCPUFE alert tests
  - interval: 1m
    input_series:
      # No alert for FE Pod High CPU
      - series: >-
          node_namespace_pod_container:container_cpu_usage_seconds_total:sum_irate{
          pod="govuk-chat-abcdefghij-abcde",
          container="app"
          }
        values: '0+0.1x7 0.7-0.1x6'
    alert_rule_test:
      - alertname: HighPodCPUFE
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert for FE Pod High CPU
      - series: >-
          node_namespace_pod_container:container_cpu_usage_seconds_total:sum_irate{
          pod="govuk-chat-abcdefghij-abcde",
          container="app"
          }
        values: '0+0.1x15'
    alert_rule_test:
      - alertname: HighPodCPUFE
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: HighPodCPUFE
              severity: critical
              container: app
              destination: slack-chat-notifications
            exp_annotations:
              summary: Elevated level of FE Pod CPU being consumed for Chat AI
              description: >-
                The FE Pod CPU usage is above 80% for a container for more than 5 minutes.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

# HighPodCPUWorker alert tests
  - interval: 1m
    input_series:
      # No alert for Worker Pod High CPU
      - series: >-
          node_namespace_pod_container:container_cpu_usage_seconds_total:sum_irate{
          pod="govuk-chat-worker-abcdefghij-abcde",
          container="worker"
          }
        values: '0+0.1x7 0.7-0.1x6'
    alert_rule_test:
      - alertname: HighPodCPUWorker
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert for Worker Pod High CPU
      - series: >-
          node_namespace_pod_container:container_cpu_usage_seconds_total:sum_irate{
          pod="govuk-chat-worker-abcdefghij-abcde",
          container="worker"
          }
        values: '0+0.1x15'
    alert_rule_test:
      - alertname: HighPodCPUWorker
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: HighPodCPUWorker
              severity: critical
              container: worker
              destination: slack-chat-notifications
            exp_annotations:
              summary: Elevated level of Worker Pod CPU being consumed for Chat AI
              description: >-
                The Worker Pod CPU usage is above 80% for a container for more than 5 minutes.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

# HighPodMemoryFE alert tests
  - interval: 1m
    input_series:
      # No alert for FE Pod High Memory
      - series: >-
          container_memory_working_set_bytes{
          job="kubelet",
          metrics_path="/metrics/cadvisor",
          pod="govuk-chat-abcdefghij-abcde",
          container="app"
          }
        values: '1073741824+1048576x15'
    alert_rule_test:
      - alertname: HighPodMemoryFE
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert for FE Pod High Memory
      - series: >-
          container_memory_working_set_bytes{
          job="kubelet",
          metrics_path="/metrics/cadvisor",
          pod="govuk-chat-abcdefghij-abcde",
          container="app"
          }
        values: '1717986918+1048576x15'
    alert_rule_test:
      - alertname: HighPodMemoryFE
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: HighPodMemoryFE
              severity: critical
              container: app
              destination: slack-chat-notifications
            exp_annotations:
              summary: Elevated level of FE Pod Memory being consumed for Chat AI
              description: >-
                The FE Pod Memory usage is above 80% for a container for more than 5 minutes.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

  - interval: 1m
    input_series:
      # No alert for Worker Pod High Memory
      - series: >-
          container_memory_working_set_bytes{
          job="kubelet",
          metrics_path="/metrics/cadvisor",
          pod="govuk-chat-worker-abcdefghij-abcde",
          container="worker"
          }
        values: '1073741824+1048576x15'
    alert_rule_test:
      - alertname: HighPodMemoryWorker
        eval_time: 15m
        exp_alerts: []

# HighPodMemoryWorker alert tests
  - interval: 1m
    input_series:
      # Alert for Worker Pod High Memory
      - series: >-
          container_memory_working_set_bytes{
          job="kubelet",
          metrics_path="/metrics/cadvisor",
          pod="govuk-chat-worker-abcdefghij-abcde",
          container="worker"
          }
        values: '1717986918+1048576x15'
    alert_rule_test:
      - alertname: HighPodMemoryWorker
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: HighPodMemoryWorker
              severity: critical
              container: worker
              destination: slack-chat-notifications
            exp_annotations:
              summary: Elevated level of Worker Pod Memory being consumed for Chat AI
              description: >-
                The Worker Pod Memory usage is above 80% for a container for more than 5 minutes.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

# NearTokensRateLimit alert tests
  - interval: 1m
    input_series:
      # No alert for OpenAI token rate limit
      - series: >-
          govuk_chat_openai_tokens_used_percentage{endpoint="test_endpoint", model="test_model"}
        values: '50+1x15'
    alert_rule_test:
      - alertname: NearTokensRateLimit
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert for OpenAI token rate limit
      - series: >-
          govuk_chat_openai_tokens_used_percentage{
          exported_endpoint="test_endpoint",
          model="test_model"
          }
        values: '50+1x14 81'
    alert_rule_test:
      - alertname: NearTokensRateLimit
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: NearTokensRateLimit
              severity: warning
              destination: slack-chat-notifications
              model: test_model
              exported_endpoint: test_endpoint
            exp_annotations:
              summary: Exceeded 80% of OpenAI token rate limit
              description: >-
                Over 80% of tokens have been used for test_model on the
                test_endpoint endpoint.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

# NearRequestsRateLimit alert tests
  - interval: 1m
    input_series:
      # No alert for OpenAI requests rate limit
      - series: >-
          govuk_chat_openai_requests_used_percentage
        values: '50+1x15'
    alert_rule_test:
      - alertname: NearRequestsRateLimit
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert for OpenAI torequestsken rate limit
      - series: >-
          govuk_chat_openai_requests_used_percentage{
          exported_endpoint="test_endpoint",
          model="test_model"
          }
        values: '50+1x14 81'
    alert_rule_test:
      - alertname: NearRequestsRateLimit
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: NearRequestsRateLimit
              severity: warning
              destination: slack-chat-notifications
              model: test_model
              exported_endpoint: test_endpoint
            exp_annotations:
              summary: Exceeded 80% of OpenAI requests rate limit
              description: >-
                Over 80% of requests have been used for test_model on the
                test_endpoint endpoint.
              runbook_url: >-
                https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

# SidekiqQueueLength and SidekiqJobAge alert tests
  - interval: 1m
    input_series:
      # Alert sustained breach tests (baseline) SidekiqQueueLength and SidekiqJobAge
      - series: 'sidekiq_queue_backlog{job="govuk-chat-worker", queue="default"}'
        values: '0 30 60 60 60 60 60'
      - series: 'sidekiq_queue_backlog{job="govuk-chat-worker", queue="answer"}'
        values: '0 10 40 40 40 40 40'
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="default"}'
        values: '0 300 1200 1200 1200 1200 1200'
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="answer"}'
        values: '0 100 800 800 800 800 800'
    alert_rule_test:
      - eval_time: 1m
        alertname: SidekiqQueueLength
        exp_alerts: []
      - eval_time: 2m
        alertname: SidekiqQueueLength
        exp_alerts: []
      - eval_time: 7m
        alertname: SidekiqQueueLength
        exp_alerts:
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: default
            exp_annotations:
              summary: Sidekiq Queue Length High
              description: >-
                The Sidekiq queue length has been over 50 items for more than 5 minutes.
                Autoscaling should have taken place, so check if the maximum number of
                worker pods are running.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html
      - eval_time: 1m
        alertname: SidekiqJobAge
        exp_alerts: []
      - eval_time: 2m
        alertname: SidekiqJobAge
        exp_alerts: []
      - eval_time: 7m
        alertname: SidekiqJobAge
        exp_alerts:
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: default
            exp_annotations:
              summary: Sidekiq Queue Items Over 15 Minutes Old
              description: >-
                Some Sidekiq items in the queue are over 15 minutes old. Check to see
                if some items have got stuck or if there is a problem processing items
                causing the delay.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

  - interval: 1m
    input_series:
      # No alert transient spike (not sustained 5m) SidekiqQueueLength
      - series: 'sidekiq_queue_backlog{job="govuk-chat-worker", queue="default"}'
        values: '0 0 60 60 60 40 40'
    alert_rule_test:
      - eval_time: 4m
        alertname: SidekiqQueueLength
        exp_alerts: []
      - eval_time: 6m
        alertname: SidekiqQueueLength
        exp_alerts: []

  - interval: 1m
    input_series:
      # No alert transient spike (not sustained 5m) SidekiqJobAge
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="default"}'
        values: '0 0 950 950 950 800 800'
    alert_rule_test:
      - eval_time: 4m
        alertname: SidekiqJobAge
        exp_alerts: []
      - eval_time: 6m
        alertname: SidekiqJobAge
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert both queues over backlog threshold sustained => two alerts at 5m SidekiqQueueLength
      - series: 'sidekiq_queue_backlog{job="govuk-chat-worker", queue="default"}'
        values: '60 60 60 60 60 60 60'
      - series: 'sidekiq_queue_backlog{job="govuk-chat-worker", queue="answer"}'
        values: '55 55 55 55 55 55 55'
    alert_rule_test:
      - eval_time: 4m
        alertname: SidekiqQueueLength
        exp_alerts: []
      - eval_time: 5m
        alertname: SidekiqQueueLength
        exp_alerts:
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: default
            exp_annotations:
              summary: Sidekiq Queue Length High
              description: >-
                The Sidekiq queue length has been over 50 items for more than 5 minutes.
                Autoscaling should have taken place, so check if the maximum number of
                worker pods are running.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: answer
            exp_annotations:
              summary: Sidekiq Queue Length High
              description: >-
                The Sidekiq queue length has been over 50 items for more than 5 minutes.
                Autoscaling should have taken place, so check if the maximum number of
                worker pods are running.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

  - interval: 1m
    input_series:
      # Alert one queue exactly at threshold (900) -> no alert; other >900 sustained -> one alert SidekiqJobAge
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="default"}'
        values: '900 900 900 900 900 900 900'
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="answer"}'
        values: '910 910 910 910 910 910 910'
    alert_rule_test:
      - eval_time: 4m
        alertname: SidekiqJobAge
        exp_alerts: []
      - eval_time: 5m
        alertname: SidekiqJobAge
        exp_alerts:
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: answer
            exp_annotations:
              summary: Sidekiq Queue Items Over 15 Minutes Old
              description: >-
                Some Sidekiq items in the queue are over 15 minutes old. Check to see
                if some items have got stuck or if there is a problem processing items
                causing the delay.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

  - interval: 1m
    input_series:
      # Alert both queues > 900 sustained -> two alerts SidekiqJobAge
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="default"}'
        values: '950 950 950 950 950 950 950'
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="high"}'
        values: '1000 1000 1000 1000 1000 1000 1000'
    alert_rule_test:
      - eval_time: 4m
        alertname: SidekiqJobAge
        exp_alerts: []
      - eval_time: 5m
        alertname: SidekiqJobAge
        exp_alerts:
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: default
            exp_annotations:
              summary: Sidekiq Queue Items Over 15 Minutes Old
              description: >-
                Some Sidekiq items in the queue are over 15 minutes old. Check to see
                if some items have got stuck or if there is a problem processing items
                causing the delay.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: high
            exp_annotations:
              summary: Sidekiq Queue Items Over 15 Minutes Old
              description: >-
                Some Sidekiq items in the queue are over 15 minutes old. Check to see
                if some items have got stuck or if there is a problem processing items
                causing the delay.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html

  - interval: 1m
    input_series:
      # Resolution after recovery SidekiqQueueLength
      - series: 'sidekiq_queue_backlog{job="govuk-chat-worker", queue="default"}'
        values: '60 60 60 60 60 60 30'
    alert_rule_test:
      - eval_time: 5m
        alertname: SidekiqQueueLength
        exp_alerts:
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: default
            exp_annotations:
              summary: Sidekiq Queue Length High
              description: >-
                The Sidekiq queue length has been over 50 items for more than 5 minutes.
                Autoscaling should have taken place, so check if the maximum number of worker
                pods are running.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html
      - eval_time: 6m
        alertname: SidekiqQueueLength
        exp_alerts: []

  - interval: 1m
    input_series:
      # Resolution after recovery SidekiqJobAge
      - series: 'sidekiq_queue_latency_seconds{job="govuk-chat-worker", queue="default"}'
        values: '950 950 950 950 950 950 800'
    alert_rule_test:
      - eval_time: 5m
        alertname: SidekiqJobAge
        exp_alerts:
          - exp_labels:
              severity: critical
              destination: slack-chat-notifications
              job: govuk-chat-worker
              queue: default
            exp_annotations:
              summary: Sidekiq Queue Items Over 15 Minutes Old
              description: >-
                Some Sidekiq items in the queue are over 15 minutes old. Check to see
                if some items have got stuck or if there is a problem processing items
                causing the delay.
              runbook_url: https://docs.publishing.service.gov.uk/manual/alerts/chat-ai-alerts.html
      - eval_time: 6m
        alertname: SidekiqJobAge
        exp_alerts: []
