rule_files:
  - fastly.yaml

evaluation_interval: 1m

tests:
  ##
  # Test case: High5xxRate happy path
  ##
  - interval: 1m
    input_series:
      - series: 'fastly_rt_origin_fetches_total{service_name="example"}'
        values: '0+10000x30' # 10k change per minute


      - series: 'fastly_origin_status_group_total{status_group="5xx", service_name="example"}'
        values: '0x10 0+50x10 50+150x10' # 0 + 0/min change (0%); 0 + 50/min (0.5%) change; 50 + 150/min change (1.5%)

    alert_rule_test:
      # Not alerting to begin with
      - eval_time: 0m
        alertname: High5xxRate

        exp_alerts: []

      # Not alerting at 10m because the rate isn't high enough
      - eval_time: 10m
        alertname: High5xxRate

        exp_alerts: []

      # Not alerting at 22m because the rate is high, but not for long enough
      - eval_time: 22m
        alertname: High5xxRate

        exp_alerts: []

      # Alerting at 30m because the rate is high and has been for long enough
      - eval_time: 30m
        alertname: High5xxRate

        exp_alerts:
          - exp_labels:
              severity: warning
              destination: slack-platform-engineering
              service_name: example
            exp_annotations:
              summary: "Fastly CDN service example is receiving elevated numbers of 5xx responses"
              grafana_path: "d/deqpydcv7ottse/fastly-global-traffic?orgId=1&from=now-1h&to=now&timezone=browser"
              runbook_url: "https://docs.google.com/document/d/1x7R2Sgpdeqd4kCiwPZp5RaeHJkOWOqgzckcMvn6tIbE/"
              description: >-
                The Fastly CDN service is receiving a higher proportion of 5xx responses
                from the origin than usual.

                Fastly CDN service: example
                5xx rate: 1.5%
  ##
  # Test case: High5xxRate does not fire for data.gov.uk
  ##
  - interval: 1m
    input_series:
      - series: 'fastly_rt_origin_fetches_total{service_name="environment data.gov.uk"}'
        values: '0+10000x20' # 10k change per minute


      - series: 'fastly_origin_status_group_total{status_group="5xx", service_name="environment data.gov.uk"}'
        values: '50+150x20' # 50 + 150/min change (1.5%)

    alert_rule_test:
      # Rate is high enough but not alerting because it's data.gov.uk
      - eval_time: 10m
        alertname: High5xxRate
        exp_alerts: []

  ##
  # Test case: High5xxRate_DataGovUk happy path
  ##
  - interval: 1m
    input_series:
      - series: 'fastly_rt_origin_fetches_total{service_name="environment data.gov.uk"}'
        values: '0+10000x60' # 10k change per minute


      - series: 'fastly_origin_status_group_total{status_group="5xx", service_name="environment data.gov.uk"}'
        values: '0x10 0+50x10 50+1500x40' # 0 + 0/min change (0%); 0 + 50/min (0.5%) change; 50 + 1000/min change (10.5%)

    alert_rule_test:
      # Not alerting to begin with
      - eval_time: 0m
        alertname: High5xxRate_DataGovUk

        exp_alerts: []

      # Not alerting at 10m because the rate isn't high enough
      - eval_time: 10m
        alertname: High5xxRate_DataGovUk

        exp_alerts: []

      # Not alerting at 23m because the rate is high, but not for long enough
      - eval_time: 23m
        alertname: High5xxRate_DataGovUk

        exp_alerts: []

      # Alerting at 55m because the rate is high and has been for long enough
      - eval_time: 55m
        alertname: High5xxRate_DataGovUk

        exp_alerts:
          - exp_labels:
              severity: warning
              destination: slack-platform-engineering
              service_name: environment data.gov.uk
            exp_annotations:
              summary: "Fastly CDN service environment data.gov.uk is receiving elevated numbers of 5xx responses"
              grafana_path: "d/deqpydcv7ottse/fastly-global-traffic?orgId=1&from=now-1h&to=now&timezone=browser"
              runbook_url: "https://docs.google.com/document/d/1x7R2Sgpdeqd4kCiwPZp5RaeHJkOWOqgzckcMvn6tIbE/"
              description: >-
                The Fastly CDN service is receiving a higher proportion of 5xx responses
                from the origin than usual.

                Fastly CDN service: environment data.gov.uk
                5xx rate: 15%

  ##
  # Test case: High5xxRate_DataGovUk does not fire for data.gov.uk
  ##
  - interval: 1m
    input_series:
      - series: 'fastly_rt_origin_fetches_total{service_name="example"}'
        values: '0+10000x30' # 10k change per minute


      - series: 'fastly_origin_status_group_total{status_group="5xx", service_name="example"}'
        values: '1000+1000x30' # 1000 + 100/min change (1%)

    alert_rule_test:
      # Rate is high enough but not alerting because it's data.gov.uk
      - eval_time: 25m
        alertname: High5xxRate_DataGovUk
        exp_alerts: []
