rule_files:
  - search_api_cronjobs.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # No alert for cronjob failed as 2 days not reached
      - series: >-
          kube_cronjob_status_last_successful_time{
          cronjob="search-api-load-page-traffic"
          }
        values: '179300+60x15'
    alert_rule_test:
      - alertname: CronjobFailed => CronjobSuccessThreshold
        eval_time: 15m
        exp_alerts: []

  - interval: 1m
    input_series:
      # Alert for cronjob failed
      - series: >-
          kube_cronjob_status_last_successful_time{
          cronjob="search-api-load-page-traffic"
          }
        values: '180300+60x15'
    alert_rule_test:
      - alertname: CronjobFailed => CronjobSuccessThreshold
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              alertname: CronjobFailed
              destination: slack-cronjob-notifications
              cronjob: search-api-load-page-traffic
              severity: critical
            exp_annotations:
              summary: search-api-load-page-traffic cronjob failing
              description: >-
                The cronjob "search-api-load-page-traffic" has not run successfully
                since 1970-01-03 02:15:00 +0000 UTC
              cronjob_uri: >-
                applications/search-api?orphaned=false&resource=&node=batch%2FCronJob%2Fapps%2Fsearch-api-load-page-traffic%2F0&tab=logs
